Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel $\rightarrow$ Restart) and then **run all cells** (in the menubar, select Cell $\rightarrow$ Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = ""
COLLABORATORS = ""
```

---

# Problem Set for L01: Git
You will perform some basic Git operations inside of this module. Make sure to validate with `nbgrader validate` or click `validate` at the top of JupyterLab's UI. You ***DO NOT*** need to be a member of the [`gt-ospo`](https://github.com/gt-ospo) GitHub organization. Additionally, make sure your local environment has Git set up for ***using SSH***. Run the following code snippets to setup the environment:

```python
!git config --global init.defaultBranch main
```

```python
import os
from pathlib import Path
from check_helper import compareOutToString, gitHashCompare, changeDirectory
import subprocess
assignmentPath = Path.home() / "PS_OSPO"
# assignmentPath = Path.cwd() # Instructor Testing Use Only
changeDirectory(assignmentPath)
```

```python
firstLastName = "JEFF" # SET THIS; This will match YOURNAME in the branch names below
```

The snippet below will **reset** your progress to test for validity during grading. Do **not** remove this code.

```python
subprocess.run(f"git push origin --delete {firstLastName}-ps1", shell=True, capture_output=True, text=True)
subprocess.run(f"git push origin --delete {firstLastName}-ps1ext", shell=True, capture_output=True, text=True)
```

From now on, use `subprocess.run()` or bash magics `!` when writing the shell commands when `%%bash` is not present at the start of a code snippet. First, ensure you have have GitHub with SSH set up (do in a jupyter terminal). Then, configure your git user name and email:

<!-- grade_id: gitconfig, points: 0, grade: false, solution: true -->
```python
%%bash
# YOUR CODE HERE
raise NotImplementedError()
```

Clone the repository at: [`gt-ospo/oss-training-templates`](https://github.com/gt-ospo/oss-training-templates) using `git clone` *with* SSH. Switch to the `newoss-nbgrader` branch. Ignore any errors if the repo already exists locally

<!-- grade_id: gitclone, points: 0, grade: false, solution: true -->
```python
%%bash
# YOUR CODE HERE
raise NotImplementedError()
git branch -a # Show all (remote) branches, not necessary for completion but visually helpful
```

<!-- grade_id: gitdirectoryexist, points: 1, grade: true, solution: false -->
```python
# Test Cases
changeDirectory(assignmentPath / "oss-training-templates")
assert (Path.cwd() / ".git").is_dir()
assert (Path.cwd() / "bubblesortfinal.py").is_file()
```

Now, create a new branch called `YOURNAME-ps1` (Replace `YOURNAME` with your name). Make sure your name has been set in the 3rd code snippet. Then, switch to this new local branch.

<!-- grade_id: cell-15761e684e80eddf, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: gitbranchexist, points: 1, grade: true, solution: false -->
```python
# Test Cases
compareOutToString("git branch --show-current", f"{firstLastName}-ps1", "Not on correct branch")
```

Add a file called `bubblesort.py`. Put the contents of `bubblesortfinal.py` into `bubblesort.py` below.

<!-- grade_id: cell-37704361e9250487, points: 0, grade: false, solution: true -->
```python
%%writefile bubblesort.py
# YOUR CODE HERE
raise NotImplementedError()
```

Commit and push the new branch from `origin/newoss-nbgrader`. **Use bash magics `!` or `subprocess.run()`**

<!-- grade_id: cell-f6247127c8c76aa3, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: gitpushremote, points: 1, grade: true, solution: false -->
```python
# Test Cases
compareOutToString("git branch -a", f"* {firstLastName}-ps1", "Branch does not exist on remote")
```

Now, create and publish a new branch based off of origin/main called `YOURNAME-ps1ext`. Make some small changes to the repo. **Use bash magics `!` or `subprocess.run()`**
1. Switch to main branch
1. Create new branch from main branch
1. Edit `README.md` using bash magics `!` or `subprocess.run()` such that it has: `'New Readme content'` as its text.
1. Push branch to remote

<!-- grade_id: cell-1360dd1fbb575ce1, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-7ccbad7a4326181f, points: 2, grade: true, solution: false -->
```python
# Test Cases
compareOutToString("git branch --show-current", f"{firstLastName}-ps1ext", "Not on correct branch")
compareOutToString("git branch -a", f"* {firstLastName}-ps1ext", "Branch does not exist on remote")
```

Now, merge branch `YOURNAME-ps1ext` into `YOURNAME-ps1`. Resolve the merge conflict by preferring the README change you just made and make sure to push changes to remote.

<!-- grade_id: cell-b333a34328d0bb47, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: gitmerge, points: 1, grade: true, solution: false -->
```python
# Test Cases
compareOutToString("cat README.md", f"New Readme content", "Incorrect README content post-merge")
```