Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel $\rightarrow$ Restart) and then **run all cells** (in the menubar, select Cell $\rightarrow$ Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = ""
COLLABORATORS = ""
```

---

# Problem Set for L02: Git

> Tip: It is recommended that you run/validate from the beginning of this notebook to check your answers rather than conditionally resetting progress through your own solution. This will save much time and effort.

You will perform some basic Git operations inside of this module. Make sure to validate with `nbgrader validate` or click `validate` at the top of JupyterLab's UI. You ***DO NOT*** need to be a member of the [`gt-ospo`](https://github.com/gt-ospo) GitHub organization. Run the following code snippets to setup the environment:

```python
# Only run this snippet once before all other snippets, do NOT run multiple times
import os
from pathlib import Path
assignmentPath = Path.cwd() # If you get an error like "file not found", make sure you are currently in the correct folder in the UI
```

```python
os.chdir(assignmentPath)
print(assignmentPath) # Make sure this matches the directory of this notebook, otherwise, restart kernel
from check_helper import compareOutToString, gitHashCompare, changeDirectory
import subprocess
```

The snippet below will **reset** your progress to test for validity during grading. Do **not** remove this code.

```python
!rm -rf oss-training-templates/
```

**Use bash magics `!` or `subprocess.run()`** when writing the shell commands when `%%bash` is not present at the start of a code snippet.

Clone the repository at: [`gt-ospo/oss-training-templates`](https://github.com/gt-ospo/oss-training-templates) using `git clone` with *either* HTTPS (recommended) or SSH. Call `changeDirectory()` into the Git repo (see test case). Then, switch to the existing `newoss-nbgrader` branch. **Ignore any errors if the repo already exists locally**

<!-- grade_id: gitclone, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE 
# Remove NotImplementedError() when finished 
raise NotImplementedError()

# Show all branches, not necessary for completion but visually helpful
changeDirectory(assignmentPath / "oss-training-templates")
!git branch -a
```

<!-- grade_id: gitdirectoryexist, points: 1, grade: true, solution: false -->
```python
# Test Cases - Do not modify or delete
changeDirectory(assignmentPath / "oss-training-templates")
assert (Path.cwd() / ".git").is_dir()
print("Success!")
assert (Path.cwd() / "bubblesortfinal.py").is_file()
print("Success!")
```

```python
# Set up default git metadata
!git config --global init.defaultBranch main
!git config user.name "Git Tutorial"
!git config user.email "git@tutorial.com"
```

Now, create a new local branch called `improved-sort`. Then, switch to this new local branch.

<!-- grade_id: cell-15761e684e80eddf, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE 
# Remove NotImplementedError() when finished 
raise NotImplementedError()
```

<!-- grade_id: gitbranchexist, points: 1, grade: true, solution: false -->
```python
# Test Cases - Do not modify or delete
compareOutToString("git branch --show-current", "Not on correct branch", f"improved-sort")
```

Add a file called `bubblesort.py`. Put the contents of `oss-training-templates/bubblesortfinal.py` into `bubblesort.py` below.

<!-- grade_id: cell-37704361e9250487, points: 0, grade: false, solution: true -->
```python
%%writefile bubblesort.py
# YOUR CODE HERE 
# Remove NotImplementedError() when finished 
raise NotImplementedError()
```

Make a commit on the new branch with the new file `bubblesort.py`. **Use bash magics `!` or `subprocess.run()`**

<!-- grade_id: cell-f6247127c8c76aa3, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE 
# Remove NotImplementedError() when finished 
raise NotImplementedError()
```

<!-- grade_id: gitpushremote, points: 1, grade: true, solution: false -->
```python
# Test Cases - Do not modify or delete
compareOutToString("git branch", "Branch does not exist locally", f"* improved-sort")
```

Now, create a new branch based off of `main` called `new-README`. **Use bash magics `!` or `subprocess.run()`**
1. Switch to main branch
1. Create new branch from main branch and switch to the new branch
1. Edit `README.md` using bash magics `!` or `subprocess.run()` such that it contains the following line:
```Markdown
This is a new README with improved information.
```

<!-- grade_id: cell-1360dd1fbb575ce1, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE 
# Remove NotImplementedError() when finished 
raise NotImplementedError()
```

<!-- grade_id: cell-7ccbad7a4326181f, points: 2, grade: true, solution: false -->
```python
# Test Cases - Do not modify or delete
compareOutToString("git branch --show-current", "Not on correct branch", f"new-README")
compareOutToString("git branch", "Branch does not exist locally", f"* new-README")
```

Now, merge branch `new-README` into `improved-sort` with the `git merge --no-edit`. If a merge conflict occurs, preferring the README change you just made.

<!-- grade_id: cell-b333a34328d0bb47, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE 
# Remove NotImplementedError() when finished 
raise NotImplementedError()
```

<!-- grade_id: gitmerge, points: 1, grade: true, solution: false -->
```python
# Test Cases - Do not modify or delete
compareOutToString("git branch --show-current", "Not on correct branch", f"improved-sort")
compareOutToString("cat README.md", "Incorrect README content post-merge", f"This is a new README with improved information.")
```

On the `improved-sort` branch, create a new file using `touch` called `OtherBranchFile.txt`. Stage this new file. Then, `stash` the file with the commit message `Move other branch file`. At the end, switch to branch `main`.

<!-- grade_id: cell-e1ca7b4d12b54302, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE 
# Remove NotImplementedError() when finished 
raise NotImplementedError()
```

<!-- grade_id: cell-b48536c5b078e397, points: 1, grade: true, solution: false -->
```python
# Test Cases - Do not modify or delete
compareOutToString("git branch --show-current", "Not on branch 'main'", f"main")
assert not os.path.exists("OtherBranchFile.txt")
compareOutToString("git stash show --name-only stash@{0}", "Stash not present or misnamed", f"OtherBranchFile.txt")
```

Now, `pop` the entry off of the stash.

<!-- grade_id: cell-6ace0fc62fe81e81, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE 
# Remove NotImplementedError() when finished 
raise NotImplementedError()
```

<!-- grade_id: cell-c457ea8fb3f65484, points: 1, grade: true, solution: false -->
```python
# Test Cases - Do not modify or delete
assert os.path.exists("OtherBranchFile.txt"), "OtherBranchFile.txt does not exist"

stash_result = subprocess.run("git stash list", shell=True, capture_output=True, text=True)
assert not stash_result.stdout.strip(), "Stash is not empty"
    
print("Success!")
```