Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel $\rightarrow$ Restart) and then **run all cells** (in the menubar, select Cell $\rightarrow$ Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = ""
COLLABORATORS = ""
```

---

# Problem Set for L03: Pytest
Insert Credits Here for flags code

You will create unit tests using pytest on http requests. Make sure to validate with `nbgrader validate` or click validate at the top of JupyterHub's UI. You will need to be a member of gt-ospo.

**Note to students and instructors**: make sure to run the following to ensure temporary files are removed before releasing assignment.

```python
!rm -f student_file.py test_unittest_hidden.py
```

Run the following segment to setup the environment:

```python
import os
from pathlib import Path
from check_helper import compareOutToString, changeDirectory
assignmentPath = Path.cwd()
!pip install pytest tqdm httpx httpcore pytest-mock
```

## Fixtures
Take a look at `flag2_common.py`, `flags2_sequential.py`, and `flags2_threadpool.py`. We will primarily focus on the threadpool implementation of `download_many`.

First, complete the two fixtures marked with "To Do" below.

<!-- grade_id: fixtures-answer, points: 0, grade: false, solution: true -->
```python
%%writefile student_file.py
import pytest
import string
import threading
from http.server import HTTPServer, SimpleHTTPRequestHandler

@pytest.fixture
def countryCodeSelLimited():
    """source/NewIntro/

    To Do: Return all possible two character combinations
    for the first 6 letters of the alphabet in a list
    """
    countryCodeList = []
    # YOUR CODE HERE
    raise NotImplementedError()
    return countryCodeList

@pytest.fixture
def countryCodeSelAll():
    """
    To Do: Return all possible two character combinations
    for all letters of the alphabet in an alphabetically ordered list
    """
    countryCodeList = []
    # YOUR CODE HERE
    raise NotImplementedError()
    return countryCodeList

@pytest.fixture
def localServerSetup():
    """
    To Do: Return code that sets up and tears down an HTTP server with host and port as listed below
    """
    host = "127.0.0.1" # do not change
    port = 8000 # do not change
    server = HTTPServer((host, port), SimpleHTTPRequestHandler)
    thread = threading.Thread(target=server.serve_forever)
    thread.daemon = True
    thread.start()

    # To Do: Return the "host:port" URL and proceed with server termination 
    # YOUR CODE HERE
    raise NotImplementedError()

```

<!-- grade_id: fixtures-test, points: 3, grade: true, solution: false -->
```python
# Test Cases - Hidden
```

Now, write a test case for the `download_one` function in `flags2_sequential`. Assuming the given inputs, test for the expected output of `download_one`. `badCC` is to produce an intentional exception.

<!-- grade_id: unittest-answer, points: 0, grade: false, solution: true -->
```python
%%writefile student_file.py
import pytest
from flags2_common import SERVERS, DEST_DIR
from flags2_sequential import download_one
from pathlib import Path

@pytest.fixture
def debugURL():
    return SERVERS['LOCAL']

@pytest.fixture
def productionURL():
    return SERVERS['REMOTE']

@pytest.mark.parametrize("cc", ["aa", "de"])
def test_download_one_basic(cc, productionURL):
    download_one(cc, productionURL)
    # To Do: Check that download_one produced a file in the correct directory
    # Note: Check destination dir in save_flag method in flags2_common
    # YOUR CODE HERE
    raise NotImplementedError()

```

<!-- grade_id: unittest-test, points: 1, grade: true, solution: false -->
```python
# Test Cases - Hidden
```

## Parameterization
Now we will take a look at some alternative backends for implementing concurrent downloads. To do so, you will implement a parameterized variant of `download_many` from `flags2_sequential.py`, `flags2_threadpool.py`, and `flags2_asyncio.py`. Fill in the parametrize section so that each backend is separately tested.

<!-- grade_id: parameterization-answer, points: 0, grade: false, solution: true -->
```python
%%writefile student_file.py
import pytest
import httpx
import string
from flags2_asyncio import download_many_asyncio
from flags2_threadpool import download_many_threadpool
from flags2_sequential import download_one, download_many_sequential
from flags2_common import main, DownloadStatus, expand_cc_args, SERVERS, DEST_DIR

@pytest.fixture
def productionURL():
    return SERVERS['REMOTE']

@pytest.fixture
def basicMultipleCC():
    return ["de", "gb"]

@pytest.fixture
def advancedMultipleCC():
    return ["ad", "ae", "al", "bg", "jp", "tv"]


# Simply test if the methods all run without error, no assertions required
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: parameterization-test, points: 1, grade: true, solution: false -->
```python
# Test Cases - Hidden
```

## Mocking

Now, implement the todos below. The following code is intended to mock the Note that `get_flag` will always be called, but `save_flag` will only be called if the flag is found in the correct local directory.

<!-- grade_id: cell-c924efbbe8115da3, points: 0, grade: false, solution: true -->
```python
%%writefile student_file.py
import pytest
import flags2_sequential
from flags2_sequential import download_many_sequential

# Note: function will be tested with ccList: ["aa", "ab", "ac", "ad"] where "ad" is the only country code with an existing flag
def download_many_intermediate(mocker, ccList, debugURL):
    numFlags = len(ccList)

    # To Do: Mock the get_flag and save_flag functions
    # To Do: Call download_many_sequential using all available parameters
    # To Do: Assert save_flag is called exactly once (see note above function header)
    # To Do: Assert get_flag is called every time
    
    # YOUR CODE HERE
    raise NotImplementedError()
    
```

<!-- grade_id: cell-f774d7b8140da27f, points: 2, grade: true, solution: false -->
```python
# Test Cases
import subprocess
import sys
contents = """import pytest
import os
from student_file import download_many_intermediate
from flags2_sequential import download_many_sequential

def test_download_many_get_flag(mocker):
    # 1
    try:
        download_many_intermediate(mocker, ["aa", "ab", "ac", "ad", "ah"], None)
    except:
        raise AssertionError("get_flag count incorrect")
    else:
        print("Success: get_flag count checked")

def test_download_many_save_flag(mocker):
    # 2
    try:
        download_many_intermediate(mocker, ["aa", "ab", "ac", "ad", "ae"], None)
    except AssertionError:
        print("Success: save_flag count checked")
    else:
        raise AssertionError("save_flag count incorrect")
"""

with open("test_unittest_hidden.py", "w") as fopen:
    fopen.write(contents)

pytestResult = subprocess.run([sys.executable, "-m","pytest", "-q", "test_unittest_hidden.py"], check=False)
assert pytestResult.returncode == 0
```