Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel $\rightarrow$ Restart) and then **run all cells** (in the menubar, select Cell $\rightarrow$ Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = ""
COLLABORATORS = ""
```

---

# Problem Set for L04: GitHub Actions

Run the following to setup the necessary dependencies:

```python
!pip install pyyaml
import yaml
!rm -f gh_config.yaml
```

In this problem set, you will create a hypothetical yaml configuration file to run unit tests for the entire workspace. The following yaml file will create a basic Github Actions configuration file. Test for:
- Python 3.10 to 3.13 (inclusive): call this `python-version`
- Ubuntu 22.04, 24.04, and latest non-LTS release: call this `os`
- Do **not** test for the combinations: `Ubuntu 22.04` & `python3.12`, `Ubuntu 22.04` & `python3.13`
- Do **not** hardcode included combinations. You should hardcode excluded combinations.

Additionally, implement two *steps* named `Install Pyright` and `Pyright General` which will install the pyright linter and then run it with default settings respectively. Use `pip` as the package manager. At the end, run pytest silently on a file called `testfile.py` inside a *step* called `Pytest Simplified`.

Note: Avoid using literal `|` or folded `>` blocks in your solutions.

<!-- grade_id: answer-yaml, points: 0, grade: false, solution: true -->
```python
%%writefile gh_config.yaml
name: Grade Project Config - PS Version

on:
  push:
    branches:
      - 'main'

jobs:
  pstest:
    strategy:
      matrix:
        # YOUR CODE HERE
        raise NotImplementedError()

    steps:
      - uses: actions/checkout@v4

      - name: Set up Python
        uses: actions/setup-python@v5
        with:
          python-version: ${{ matrix.python-version }}

      - name: Display Python version
        run: python3 --version

      - name: Install dependencies
        run: |
          python3 -m pip install --upgrade pip
          pip install -r requirements.txt

        # YOUR CODE HERE
        raise NotImplementedError()

```

<!-- grade_id: test-yaml, points: 2, grade: true, solution: false -->
```python
# Test Cases - Hidden
from collections import defaultdict
with open("gh_config.yaml", "r", encoding="utf-8") as fopen:
    data = yaml.safe_load(fopen)

```

<!-- grade_id: test-yaml-pyright, points: 1, grade: true, solution: false -->
```python
# Test Cases
# Check for correct pyright configuration
visited = 0
psConfig = data.get("jobs").get("pstest")
for cRule in psConfig.get("steps"):
    cRun = str(cRule.get('run', ''))
    cName = str(cRule.get('name', '')).strip()
    if ("Install Pyright" in cName):
        visited += 1
        assert "pip install" in cRun and "pyright" in cRun, "Invalid pyright installation"
    elif ("General Pyright" in cName):
        visited += 1
        assert "pyright" in cRun, "Pyright configuration incorrect"

assert visited == 2, "Pyright configuration incorrect"
```

<!-- grade_id: test-unittest, points: 2, grade: true, solution: false -->
```python
# Test Cases - Hidden
```