Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel $\rightarrow$ Restart) and then **run all cells** (in the menubar, select Cell $\rightarrow$ Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = ""
COLLABORATORS = ""
```

---

# Problem Set for L10: Python Packaging

## Instructions

If *multiple* answers are correct (there will always be between 1 and 4 correct answers inclusive), separate answers with a comma in between. Use the letters associated with the answer to denote selection on the line **after**: `%%writefile -a ps10_student_answers.txt`. Answers should contain comma-separated letters e.g. `A, B` if there are multiple answers or just the letter e.g. `A` for one answer. The order of answers does not matter. 

Please utilize the reference material at [intersect training](https://intersect-training.org/packaging/index.html). The content follows sections 2, 4, 6, 8, 9.

> Tip: It is recommended that you run/validate from the beginning of this notebook to check your answers rather than conditionally resetting progress through your own solution. This will save much time and effort.

```python
# Only run this snippet once before all other snippets, do NOT run multiple times
import os
from pathlib import Path
assignmentPath = Path.cwd() # If you get an error like "file not found", make sure you are currently in the correct folder in the UI
```

```python
os.chdir(assignmentPath)
print(assignmentPath) # Make sure this matches the directory of this notebook, otherwise, restart kernel
```

```python
%%writefile ps10_student_answers.txt
Student Solutions:
```

#### 1. To create a package with a method "yourfunction", what files would be required? Assume you are in the root directory of your package.
A. `./src/yourfunction.py`  
B. `./src/main.py`  
C. `./src/\_\_init\_\_.py`   
D. `./src/pyproject.toml`  
E. `./pyproject.toml`  

<!-- grade_id: cell-63af2fac119193ea, points: 0, grade: false, solution: true -->
```python
%%writefile -a ps10_student_answers.txt
# YOUR CODE HERE 
# Remove NotImplementedError() when finished 
raise NotImplementedError()
```

#### 2. Which lines of metadata in `pyproject.toml` are ***correct*** for a package if it requires at least Python 3.5, PySide6 6.10.3, and uses setuptools for its build. Assume the author is called `Some Author` with email `someauthor@example.com`. Assume this software is licensed under LGPLv3. Note that the backend for setuptools is `setuptools.build_meta`.

A. Under `[project]`: `requires-python = ">=3.5"`  
B. Under `[project]`: `dependencies = [ "PySide6" ]`  
C. Under `[build-system]`: `requires = "setuptools"`  
D. Under `[project]`: `authors = [ name = "Some Author", email = "someauthor@example.com" ]`  
E. Under `[project]`: `classifiers = ["License :: OSI Approved :: LGPLv3 License"]`  

<!-- grade_id: cell-a904f2991a8cb6a2, points: 0, grade: false, solution: true -->
```python
%%writefile -a ps10_student_answers.txt
# YOUR CODE HERE 
# Remove NotImplementedError() when finished 
raise NotImplementedError()
```

#### 3. Which of the following statements about pip and pipx are correct?

A. `pip` installs packages into the system environment  
B. `pipx` installs packages into a virtual environment  
C. `pip` always makes a package visible to your **terminal** environment  
D. `pipx` caches a virtual environment upon each package run  
E. `pipx` is included with most Python installations 

<!-- grade_id: cell-15b456cec662563a, points: 0, grade: false, solution: true -->
```python
%%writefile -a ps10_student_answers.txt
# YOUR CODE HERE 
# Remove NotImplementedError() when finished 
raise NotImplementedError()
```

#### 4. What is/are the difference(s) between `project.dependencies` and `build-system.requires`?

A. `build-system.requires` lists runtime dependencies  
B. `build-system.requires` is used to build wheels    
C. `project.dependencies` lists runtime dependencies  
D. `build-system.requires` packages will be installed on the end-user's machine  
E. `project.dependencies` packages will be installed on the project developer's machine     

<!-- grade_id: cell-229081f0d8676d21, points: 0, grade: false, solution: true -->
```python
%%writefile -a ps10_student_answers.txt
# YOUR CODE HERE 
# Remove NotImplementedError() when finished 
raise NotImplementedError()
```

#### 5. Which of the folowing are true of sdists and wheels?

A. Wheels are always platform independent  
B. sdists contain prebuilt binaries of the package   
C. Wheels tend to install faster (in less time)  
D. Wheels strip `pyproject.toml` and other metadata from the package  
E. sdists typically require packages listed in `build-system.requires`  

<!-- grade_id: cell-a9e2958859e10b8b, points: 0, grade: false, solution: true -->
```python
%%writefile -a ps10_student_answers.txt
# YOUR CODE HERE 
# Remove NotImplementedError() when finished 
raise NotImplementedError()
```

#### 6. Which statements about setting up Mkdocs or Sphinx are correct?

A. In Mkdocs: Documentation (markdown) files belong in the project's root directory  
B. In Mkdocs: Specific plugins can be used to autogenerate API documentation from code  
C. In Mkdocs: `index.md` contains documentation found in the repository README  
D. In Sphinx: Documentation (markdown) files belong in the `/docs` directory  
E. In Sphinx: Dependencies (including sphinx itself) used for documentation belong in `project.dependencies`  

<!-- grade_id: cell-e3c1ec606caf29ac, points: 0, grade: false, solution: true -->
```python
%%writefile -a ps10_student_answers.txt
# YOUR CODE HERE 
# Remove NotImplementedError() when finished 
raise NotImplementedError()
```

Make sure to validate your submission to see if `ps10_student_answers.txt` has the desired answers in order.

<!-- grade_id: cell-e3bb9c3ae0c2a8e6, points: 6, grade: true, solution: false -->
```python
# Test Cases (Hidden) - Do not modify or delete
```

Now, fill in the blanks to create a `pyproject.toml` file that satisfies the following requirements:

1. The project name is `PS10`, the version is `1.1.1`
1. The project has mandatory dependencies `scipy` and `matplotlib`. It also has optional dependency `mkdocs-material` which should be under the name `docs`.
1. The project should be licensed under the `FSF Approved` `GPLv2 License` (Use the wording exactly as seen)
1. The project cannot be uploaded to PYPI using normal means.

Anything not explictly request by the above statements does not need to be added.

<!-- grade_id: cell-1224bcaa90004f5f, points: 0, grade: false, solution: true -->
```python
%%writefile pyproject.toml
[build-system]
requires = ["hatchling"]
build-backend = "hatchling.build"

[project]
# YOUR CODE HERE 
# Remove NotImplementedError() when finished 
raise NotImplementedError()
authors = [
  { name="Student", email="student@example.com" },
]
description = "A sample pyproject.toml test"
readme = "README.md"
requires-python = ">=3.14"
# YOUR CODE HERE 
# Remove NotImplementedError() when finished 
raise NotImplementedError()
```

<!-- grade_id: cell-d27e4c8575fdcf91, points: 4, grade: true, solution: false -->
```python
import tomllib

ANSWER = {
    "project": {
        "name": "PS10",
        "version": "1.1.1",
        "dependencies": ["scipy", "matplotlib"],
        "classifiers": [
            "License :: FSF Approved :: GPLv2 License",
            "Private :: Do Not Upload",
        ],
        "optional-dependencies": {
            "docs": ["mkdocs-material"]
        },
    }
}

def verify_subset(actual, expected):
    if isinstance(expected, dict):
        return (
            isinstance(actual, dict)
            and all(
                (key in actual) and verify_subset(actual[key], value)
                for key, value in expected.items()
            )
        )

    if isinstance(expected, list):
        if not isinstance(actual, list):
            return False

        return all(
            any(verify_subset(a, e) for a in actual)
            for e in expected
        )

    return actual == expected # Default

with open("pyproject.toml", "rb") as fopen:
    pyproject = tomllib.load(fopen)

assert verify_subset(pyproject, ANSWER), "Pyproject.toml does not match specifications"
print("Success!")
```