Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel $\rightarrow$ Restart) and then **run all cells** (in the menubar, select Cell $\rightarrow$ Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = ""
COLLABORATORS = ""
```

---

# Module 01 - Git Basics

<img src="./img/icon-git.png" width="200" />

# Contents
- [1: About this training module](#1:-About-this-training-module)
- [2: Git and GitHub Setup](#2:-Git-and-GitHub-setup)
- [3: About Git](#3:-About-Git)
- [4: Working with a Git repo](#4:-Working-with-a-Git-repo)
  * [4.1: Creating a repo and committing files](#4.1:-Creating-a-repo-and-committing-files)
  * [4.2: Working with remotes](#4.2:-Working-with-remotes)
  * [4.3: Downloading a repo](#4.3:-Downloading-a-repo)
  * [4.4: Adding more changes](#4.4:-Adding-more-changes)
  * [4.5: Getting new changes from the remote](#4.5:-Getting-new-changes-from-the-remote)
- [5: Branches](#5:-Branches)
  * [5.1: About Git branches](#5.1:-About-Git-branches)
  * [5.2: Checking-out changes](#5.2:-Checking-out-changes)
  * [5.3: Using topic branches](#5.3:-Using-topic-branches)
  * [5.4: Merging branches](#5.4:-Merging-branches)
- [6: Git Commit Signing and DCO Signoff](#6:-Git-Commit-Signing-and-DCO-Signoff)
  * [6.1: Developer Certificate of Origin (DCO) Signoff](#6.1:-Developer-Certificate-of-Origin-(DCO)-Signoff)
  * [6.2: Git Commit Signing (Cryptographic Signatures)](#6.2:-Git-Commit-Signing-(Cryptographic-Signatures))

# 1: About this training module

Welcome to the GT OSPO VSIP!  This training module will introduce you to concepts of Git, especially related to branching and merging.  Understanding Git branching/merging will be essential to your success in as code contributor.  

### How should I use this training module?

Ensure all sections are completed correctly by either using the "validate" function of this notebook or by restarting the kernel and running the notebook from top to bottom. You'll get the most out of this module if you run the code snippets, but it is ***completely optional*** and skimming is fine. Bash sections are denoted with `%%bash` and bash oneliners (magics) are prepended with `!`. These belong in code cells along with other Python code.


Run the following to setup the python environment for the remainder of the module. Not doing so will result in failed tests:

```python
import os
import subprocess
from pathlib import Path
from check_helper import compareOutToString, gitHashCompare, changeDirectory
assignmentPath = Path.cwd()
```

### Why aren't we learning to use the Git GUI in (VS Code, PyCharm, ...)?

Good question!  We love our IDEs, we use Git in our IDEs all the time, and so should you.  **But** it really helps to know how to precisely use Git in the command line, especially if your IDE isn't doing exactly what you want.  Trust us, it can really help you out of a jam.  

Another issue is that not everyone uses the same IDE.  All the terms and concepts you learn here will translate directly to whatever IDE you use.

# 2: Git and GitHub setup

<img src="./img/icon-gears.png" width="200" />

## Installing Git

GitHub has excellent instructions for installing Git on macOS, Windows, and Linux:  https://github.com/git-guides/install-git

If you are on Windows, be sure include the 'Git Bash' terminal during the installation (it should be selected by default).  You can use Git Bash to follow-along with this tutorial.  If you are on macOS or Linux, you can use any terminal to follow-along.

## First-time Git Setup

After installing Git, you should set your username an email.  These will be used to identify changes that you've made to a Git repo.  In the terminal run `git config` with your own name email:


```python
%%bash
git config --global user.name 'Your Name'
git config --global user.email 'youremail@provider.com'
git config --global init.defaultBranch main
```

The `user.name` is arbitrary and is only for the benefit of other humans.  However, the `user.email` is quite important when you're working with GitHub, GitLab, or another service.  This email is used to link your GitHub account to any commits in a Git repo.  It's not a bad idea to use a `user.email` that you'll have access to forever, instead of your GT email.  This will ensure that you can always keep your GitHub account linked to commits in Git repos.   

For more first-time setup options, see here:  https://git-scm.com/book/en/v2/Getting-Started-First-Time-Git-Setup

## Registering for GitHub

If you don't already have a GitHub account, go to https://github.com/ and create one.  Again, it's not a bad idea to use an email address that you'll have access to forever, instead of your GT email.  

## Setting up authentication with GitHub

When using Git, there are several method to authenticate with GitHub.  These days, username/password is *not* one of those methods!  The simplest method is via SSH.  See the GitHub docs section "Connecting to GitHub with SSH" (https://docs.github.com/en/authentication/connecting-to-github-with-ssh), which has instructions for every operating system.


# 3: About Git

<img src="./img/icon-spill.png" width="200" />


## What is Git?

Git is the most widely used **version control system (VCS)** today.  Any VCS uses a database to record changes to a set of files over time, such as the source files in a software project.  The changes over time are called the history **history**.  Using the database, you can retrieve specific **versions** of files from any point in history.

## Why is Git?

<img src="./img/it-worked-yesterday.png" width="200" />

One of Git's biggest advantages is to make **collaboration** easier.  By using a branching workflow (which we'll discuss below), multiple developers can work independently on the same codebase.  

Another advantage of any VCS is the ability to **respond to bugs**.  If your app is in production, and a bug is discovered, then you can revert the app to a previous stable while you investigate the cause.  The VCS also helps you discover which **specific change and developer** instroduced the bug.  The intent is not to blame or shame anyone.  Rather, the idea is that the developer who introduced the bug is the one who is most likely to know that part of the codebase and respond to that bug quickly.

## Parts of a Git Project Directory

The contents of a Git repo are stored on your computer as normal files in a normal directory.  In Git terminology, these files are the **working tree**.  You can edit these files freely with your usual text editors and IDEs.  

On your computer, changes to the files are stored in the **local repository** (or local **database**).  Soon, we will learn how to store (or **commit**) new versions of the files; and how to retrieve (or **checkout**) previous versions of the files.  

A small project directory is shown below. 
* The database and metadata are stored in the hidden `.git/` directory.  You can treat this as a black box, and your only interaction with the database and metadata should be via `git` commands.
* `my_program.py` and `README.md` are files you are actively working on.  You definitely want to record these in the local repo.
* `__pycache__` contains intermediate bytecode from running your Python program.  You probably do *not* want to record these in the local repo...

``` Console
~/my_git_project$ ls -la
total 95
drwxr-xr-x  4 gburdell  staff   6 Jan 25 12:33 .
drwx--x--x 43 gburdell  staff  97 Jan 25 12:33 ..
drwxr-xr-x  7 gburdell  staff  10 Jan 25 12:33 .git
-rw-r--r--  1 gburdell  staff  69 Jan 25 12:32 my_program.py
drwxr-xr-x  2 gburdell  staff   3 Jan 25 12:33 __pycache__
-rw-r--r--  1 gburdell  staff  51 Jan 25 12:33 README.md
```

## Tracked vs. Untracked Files

Git does not automatically add files from the working tree to the repo.  Instead, the user must explicitly specify which files are stored in the repo (**tracked**) and which are not (**untracked**).  It's important to consider which files to *not* track.  

In a Python project, you should not track byte code (`*.pyc` and `*.pyo` in `__pycache__` directories) since they will be automatically regenerated whenever someone else runs it with a different operating system or Python version.  Tracking them will just take up unnecessary space in the database. 

For C/C++ and other compiled languages, you should not track build artifacts (`*.o` files, libraries, executables) for a simlar reason.  These artifacts will usually not work on a different operating system.

# 4: Working with a Git repo


<img src="./img/icon-plates.png" width="200" />

In this section, we will create a local repo, add files, and work remote repo on GitHub

## 4.1: Creating a repo and committing files

### `git init`: Creating a local repo

The `git init` command creates a local repo in the current directory.  Let's create a new directory, `~/git-workflow` and initialize a repo inside it. The next section will *reset* the working directory state for every full run; ignore any errors that may appear.

(You can use a directory with a different name and location, if you like, but future code snippets will need to be modified)

```python
# This snippet resets the local state for every run
!rm -rf git-workflow
!mkdir -p ./git-workflow
!cd ./git-workflow
!git init
```

You can see the `.git/` directory was created for the database and metadata.  

```python
%%bash
cd ./git-workflow
ls -lha
```

Run the following cell once to set your directory to the new folder.

```python
changeDirectory(assignmentPath / "git-workflow")
```

<!-- grade_id: gitinit, points: 1, grade: true, solution: false -->
```python
# Test Cases
assert (Path.cwd() / ".git").is_dir()
```

### `git status`: Show the state of the working tree

`git status` shows the state of tracked files, untracked files, and the database. Right now, the working tree is empty, and the database is empty, so there's not much to report.  

We already have a **branch** called `main`, which was created when the repo was initialized.  We'll cover branches in detail soon.  For now, you can think of them like branches of a family tree that describe your project's history.  

```python
!git status
```

<!-- grade_id: gitstatusempty, points: 1, grade: true, solution: false -->
```python
# Test Cases
compareOutToString("git status", "Incorrect 'git status'", "nothing to commit")
```

### Creating new files

To create and edit files in your working tree, you can use any text editor.  In `~/git_workflow`, let's create two files.

* First, create a file named `README.md` file with these contents.  You might notice a mistake, but don't correct it right now! :)

``` text
# Git Workflow

This is an example repo for the GT OSPO VSIP Spring 2024 Program.  
```

* Then, create a Python source file named `my_abs.py` with these contents. We will be making changes to `my_abs.py` throughout this demo.  

``` python
def my_abs(x):
    if x < 0:
        return -x
    else:
        return x
```

<!-- grade_id: cell-newfileanswer, points: 0, grade: false, solution: true -->
```python
%%bash
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-newfilerealstatus, points: 2, grade: true, solution: false -->
```python
# Test Cases
assert (Path.cwd() / "my_abs.py").is_file()
assert (Path.cwd() / "README.md").is_file()
```

### Untracked files

`git status` reports that the new files are **untracked**.  Recall that untracked files are not recorded in repo's database, and that Git does not automatically track new files.  Git helpfully suggests that we should track them with the `git add` command, and we'll do that soon.  But first let's discuss the possible **states** for files in your working tree.  Knowing these states can help you troubleshoot many issues when you're working with your repo. 

```python
!git status
```

<!-- grade_id: untrackedfile1, points: 1, grade: true, solution: false -->
```python
# Test Cases
compareOutToString("git status", "Incorrect 'git status'", "untracked files present")
```

### The state of files in a working tree

Files in the working tree are either tracked or untracked:

* **Untracked**: No versions of the file are stored in the repo's database
* **Tracked**:  One or more versions of the file are in the database

Additionally, changes in tracked files are always in one of three states:

* **Unmodified**:  The file in the working is tree are up-to-date with the database.  Git sometimes refers to these changes as "up-to-date".  
* **Modfied**:  The file has changes that have not yet been stored in the database.  Furthermore, the user hasn't specified that these changes will be stored in the next database update.  
* **Staged**: The file's changes that will be stored in the next database update.  Git sometimes refers to these as "to be committed".  


<img src="./img/lifecycle.png" width="800" />


(Image credit:  Scott Chacon and Ben Straub.  Pro Git, [Section 2.2](https://git-scm.com/book/en/v2/Git-Basics-Recording-Changes-to-the-Repository))

### `git add`: Stage new changes



In our working copy of `git_workflow`, we have two untracked files: `README.md` and `my_abs.py`.  To record their changes, we need two steps (as shown in the [diagram above](#The-state-of-files-in-a-working-tree)):

1. **Add the files:**  This changes the files' states from "Untracked" to "Staged".  At this point, the changes are *ready* to be stored in the database but *are not yet stored*.  
2. **Commit**:  This changes their state from "Staged" to "Unmodified".  At this point, the changes have actually been stored in the database.

To accomplish Step 1 (Untracked ➞ Staged), we use `git add`.  Its usage is:

```
git add <file1> [<file2> ...]
```

```python
!git add README.md my_abs.py
```

Now `git status` shows the files have changes that are "to be committed".  This means the changes are staged for the next database update.  But remember that the database has **not** actually been updated.

```python
!git status
```

<!-- grade_id: gitaddbasic, points: 1, grade: true, solution: false -->
```python
# Test Cases
compareOutToString("git status", "Incorrect 'git status'", "Changes to be committed")
```

### Modifying a staged file

What happens when you've staged a file, but before you commit it, you realize you need to fix something?  For example, in `README.md`, we made a typo.  We wrote "Spring 2024" instead of "Summer 2024".  

Let's correct `README.md` with our text editor, and then look at the `git status`:

<!-- grade_id: cell-55d601e9e91ea6a8, points: 0, grade: false, solution: true -->
```python
%%bash
# YOUR CODE HERE
raise NotImplementedError()
```

```python
!git status
```

<!-- grade_id: modifystagefile, points: 1, grade: true, solution: false -->
```python
# Test Cases
compareOutToString("git status", "Incorrect 'git status'", "Changes not staged for commit")
```

Git tells us:
* `README.md` and `my_abs.py` have changes that are staged ("to be committed")
* `README.md` also has changes that are not staged

How can one file have both staged and unstaged changes?  The reason is that `git add` stages the state of the file **at the exact moment** you run `git add`.  So if you run `git add`, and make additional changes afterwards (like changing "Spring" to "Summer"), then those additional changes are not automatically staged.  

### Staging new content (again)

To fix this, we'll run `git add` again to stage the new changes ("Spring" to "Summer").  You can see that Git suggests this, too, when it says: 'use "git add <file>..." to update what will be committed'

```python
!git add README.md
```

Now we see that all the changes in `README.md` have been staged, since there are no "Changes not staged for commit" anymore.  

```python
!git status
```

### `git commit`: Updating the database

Finally, we'll use `git commit` to record the previously-staged changes to the database.  

You can use the `-m` option to specify a **commit message** on the command line.  This is a short message that lets other humans know what changes you've made.  Many projects have conventions about what info should go into a commit message.  Ask your project manager for details.  

(If you do not use `-m`, a text editor will pop up and prompt you to enter a message.)

```python
!git commit -m "First commit of my_abs (no try/except yet)"
```

### `git log`:  Showing the repo's history

Now that we actually have information in our repo's database, we can use `git log` to show the history.  `git log` has many options to show more or less information about the history.  You can run `git help log` to see all the available options.

<!-- grade_id: gitloganycommit, points: 1, grade: true, solution: false -->
```python
!git log
```

## 4.2: Working with remotes

Up to this point, we've only made changes to the **local** repo on our computer.  We haven't touched GitHub at all.  Now, we are going to create a **remote** repo on GitHub and upload our local repo to it.

### Creating a remote repo

From the front page of [GitHub](https://github.com), you can create a new repo by clicking on the the "+" button on the top menu bar.  You can also go directly to https://github.com/new

<img src="./img/working-with-remote-01.png" width="500" />

You should now see the "Create a new repository" page.  Since we are uploading an existing repository (instead of creating a new one), we'll only need a few options.

* **Repository template**:  GitHub has starter templates (with directories and some boilerplate code) for common types of projects.  We'll select "No template".
* **Owner**:  The repo's owner can be an individual account or an organization (a group of accounts).  A GitHub organization is a powerful tool for collaborating and managing permissions, and more likely than not, you'll be working in an organization for your VSIP project.  For this tutorial, let's use your individual GitHub account.
* **Repository name**: This will be the last part of the repo's URL.  When someone clones your repo, it will also be the default name for the project directory.  Let's give it the same name as your existing project directory.
* **Visibility**: You can make your repo either Public or Private, whichever you prefer
* **Initialize this repository** and below:  Since we are uploading an existing repo, we should not create a `README`, `.gitignore`, or license.

When you're finished, click the "Create repository" button

<img src="./img/working-with-remote-02.png" width="500" />

You will now be taken to your repo's webpage, and you should see the following info.  We're going to focus on the section "...or push an existing repo from the command line."  In particular, we will focus on the `git remote add` command and the `git push` command. 

(*The `git branch -M main` command is intended for legacy Git repos that would like to change their default branch name.  Historically, Git's default branch name was "master".  Starting around 2020, the Git community began efforts to change this naming convention to "main".  This was widely embraced by stakeholders in the community, such as [Git](https://lore.kernel.org/git/pull.656.v4.git.1593009996.gitgitgadget@gmail.com/) itself, [GitHub](https://github.blog/changelog/2020-10-01-the-default-branch-for-newly-created-repositories-is-now-main/), [GitLab](https://about.gitlab.com/blog/2021/03/10/new-git-default-branch-name/), and [BitBucket](https://bitbucket.org/blog/moving-away-from-master-as-the-default-name-for-branches-in-git).  Today, in recent versions of Git, the default branch name is "main", so you will not need to explicitly change the default with `git branch -M main`*)

<img src="./img/working-with-remote-03.png" width="600" />

### `git remote`:  Add and manage remotes

The `git remote add` command lets you link your local repo to a remote repo.  The syntax is:
```
git remote add <name> <url>
```
The `<url>` is the URL of your repo on GitHub.  The `<name>` is a shorter identifier for that URL, which makes life easier for humans.  By convention, humans use "origin" as the name for the default remote.  When you are collaborating with a team and have multiple remotes (which is quite common!), there are naming conventions for other remotes (such as "upstream").

Here, we only have one remote, so we name it "origin".  Make sure you use the URL from your own GitHub repo!

```python
!git remote add origin git@github.com:YourUsername/git-workflow.git
```

Now you can use the `git remote -v` command to see which remotes are defined.  You can see that this remote is setup for both fetching (downloading) and pushing (uploading).

```python
!git remote -v
```

### `git push`: Upload to the remote repo

The `git push` command is used to upload changes from your local repo to a defined remote repo.  The syntax is:

```
git push <remote> <branch>
```

A useful first-time option is `-u`.  It sets `<remote>` as the default remote for `<branch>`.  We'll use that below.  Th effect is that, in subsequent pushes, we can use `git push` without any extra arguments to push `main` to `origin`

<!-- grade_id: gitpushsuccess, points: 1, grade: true, solution: false -->
```python
!git push -u origin main
```

### Back to GitHub

Now you can go back to your repo's webpage on GitHub and see your repo's contents!  You might need to click the refresh button in your browser to see the updates.  

## 4.3: Downloading a repo

### `git clone`:  Download a repo

Now that your GitHub repo is populated, let's try downloading (or **cloning**) it somewhere else.  It's common to clone your repo on multiple computers depending on where you're working (like your laptop, a remote server, etc.).  For this demo, we'll just create another directory in `~/somewhere-else`

(You can use a different name or location, if you like)

```python
changeDirectory(assignmentPath)
!mkdir -p ./somewhere-else
changeDirectory(assignmentPath / "somewhere-else")
```

Now we can use the `git clone` command to download the repo.  The syntax is:

```
git clone <url>
```

The URL is the same one you used with `git remote add`.  This will create a working copy of the local repo in a new directory.  The directory's name is the last part of the URL.  Afterwards,  we can enter the new working copy and see that our files are present.  We can also verify that our repo's history is present.

```python
%%bash
rm -rf git-workflow
git clone git@github.com:YourUsername/git-workflow.git
```

```python
changeDirectory(Path.cwd() / "git-workflow")
```

```python
!ls -la
```

<!-- grade_id: gitclonebasic, points: 1, grade: true, solution: false -->
```python
!git log
```

## 4.4: Adding more changes

### Back to our original working copy

Let's return to our original working copy and make some more changes

```python
changeDirectory(assignmentPath / "git-workflow")
```

### Making more changes

Let's add some new features to `my_abs.py`.  Copy/paste this into my_abs.py


``` python
import math

def my_abs(x):
    try:
        if x < 0:
            return -x
        else:
            return x
    except TypeError:
        return math.nan
```

<!-- grade_id: cell-e9a5027cb0296ed5, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

### `git diff`: Show changes in modified files

At this point, `my_abs.py` is the in the "modified" state.  It has changes that have not yet been staged.  It can be very useful to see changes before you stage and commit. `git diff` lets you do exactly that.

The appearance of a "diff" is the same in many places.  GitHub and your IDE will use a very similar appearance when showing you differences in files.  Hence, understanding diffs is a key skill for understanding your code's history.  The meaning is:
* Lines that are green and/or begin with a "+" have been added to the newer version but are not present in the older vesion.
* Lines that are red and/or begin with a "-" are present in the older version but have been omitted in the newer version.
* Lines that are not colored are the same in both versions

```python
!git diff
```

<!-- grade_id: gitdiffbasic, points: 1, grade: true, solution: false -->
```python
# Test Cases
compareOutToString("git diff", "Incorrect 'git diff' - check the content of my_abs.py", ["except TypeError:", "return math.nan"], False)
```

### Staging and committing

As we mentioned, `my_abs.py` now has unstaged changes.  We can stage and commit them in one shot with the `-a` flag.  This will stage and commit changes to all currently **tracked** files.  

After committing, we can verify that there is a new commit in our history.  We'll use some additional options to show a more condensed log.  Logs are, by default, shown from most to least recent.  

```python
!git status
```

```python
!git commit -am "Added try/except"
```

```python
!git log --oneline --graph --branches --remotes
```

<!-- grade_id: gitlogimmediateoutput, points: 1, grade: true, solution: false -->
```python
# Test Cases
compareOutToString("git log --oneline --graph --branches --remotes", "Incorrect 'git log'", "(HEAD -> main) Added try/except")
```

### Pushing to remote

Finally, we'll push our updates to GitHub.  Go back to your GitHub webpage and see your updates!  (You might have to click refresh again.)

<!-- grade_id: gitpushsuccess2, points: 1, grade: true, solution: false -->
```python
!git push
```

## 4.5: Getting new changes from the remote

Like most things in Git, you have to be extremely explicit about updating your working copies.  If you push changes from one working copy, then other working copies are **not** automatically updated. You must explicitly download the new changes in the other working copies.  

### What's happening in our other working copy?

Let's go back to other working copy in `./somewhere-else/git-workflow`.  The log shows that the working copy does *not* have the changes that we just pushed.  And if we look at `my_abs.py`, we'll see it's still at the older version.  

```python
changeDirectory(assignmentPath / "somewhere-else" / "git-workflow")
```

```python
!git log
```

```python
!cat my_abs.py
```

### `git pull`: Download changes from the remote

We will bring this working copy up-to-date with `git pull`.  The syntax is:

```
git pull [<remote> <branch_name>]
```
If the branch has a default remote, then we can omit the extra arguments and just run `git pull`.  Now when we look at the log and the contents of `my_abs.py`, we see that everything's up-to-date with the remote.

```python
!git pull
```

```python
!git log
```

<!-- grade_id: gitpulldefault, points: 1, grade: true, solution: false -->
```python
# Test Cases
compareOutToString("git log", "Incorrect 'git pull'", ["Added try/except", "First commit of my_abs (no try/except yet)"])
```

```python
!cat my_abs.py
```

### Back to our first working copy

We'll be working with our first working copy in the next section, so let's make sure to move back:

```python
changeDirectory(assignmentPath / "git-workflow")
```

# 5: Branches

<img src="./img/icon-branch.png" width="200" />

At this point, we've been discussing one branch, "main".  Now we will consider additional branches, and we'll see how branches can help you in collaborative workflows.

## 5.1: About Git branches

### Git branches are pointers

It's common to think about Git branches like the branches of a family tree.  However, it's more useful to remember that **a Git branch is a pointer to a commit**.  

In older VCSs, a "branch" was implemented as a complete copy of the files.  This used a large amount of space in the database, and hence, branches were a less common part of workflows.

In contrast, Git branches are fast and cheap with very little storage or performance overhead.  They are an integral part of any Git workflow.

### Our branches

In fact, we already have multiple branches.  If you look at `git log` (make sure you've moved back to `~/git-workflow` !) you'll see:

* **HEAD** is a special branch that points to the current state of our working tree.
* **main** is the branch we created when we initialized the repo.
* **origin/main** is a remote-tracking branch.  It reflects the state of our remote repo "origin" on GitHub.  More on this soon...



<img src="./img/branches-01.png" width="400" />

## 5.2: Checking-out changes

### `git checkout` : Changing the files in the working tree

Suppose you want to work with the files from a different point in time.  We use `git checkout`, which has the syntax:

```
git checkout <branch | commit>
```

First let's checkout the files from our first commit.  Every commit is uniquely identified by an alphanumeric hash, and we can use this to specify a commit in `git checkout` and other commands.

In the repo history from this demo, our first commit has the hash ID `0565d4790146c2efaa59041453e7c7c09baa3cc1`.  You can use abbreviated hashes (as long as the abbreviation still uniquely identifies a commit), and it's common to use the first 7 characters.  So we will use `0565d47`. 

(*Your repo will almost certainly have different hashes, so as you're following along, be sure to use your repo's hash.*)

```python
!git checkout 0565d47
```

### Results of checkout

When we look at the log, we can see that the HEAD has moved back to "First commit".  Likewise, `my_abs.py` in the working tree is at the state of "First commit" -- it does not have the try/except .  However, note that the branches "main" and "origin/main" are unchanged.

```python
!git log --oneline --graph --branches --remotes 
```

<!-- grade_id: gitcheckoutnew, points: 1, grade: true, solution: false -->
```python
# Test Cases
compareOutToString("git log --oneline --graph --branches --remotes", "Incorrect 'git checkout'", ["(HEAD) First commit of my_abs (no try/except yet)"])
```

```python
!cat my_abs.py
```

<img src="./img/branches-02.png" width="400" />

### Detached HEAD

What about the scary message about the detached HEAD?  This is perfectly normal and common!  All it means is that HEAD is not pointing to the same commit as any other branch.  If you're exploring your repo's history, you will encounter this message often.  You can suppress this message with:

```
git config --global advice.detachedHead false
```

Or you can learn to ignore it :)

### Checking out a branch

You can also specify a branch name to `git commit`.  This is probably the more common usage.

Let's return our working tree to the state of "main".  You will notice that:
* The HEAD moves back to "main"
* The state of the working tree is also back to "main". Notices that `my_abs.py` has the try/except as expected.

```python
!git checkout main
```

```python
!git log --oneline --graph --branches --remotes 
```

<!-- grade_id: gitcheckoutmain, points: 1, grade: true, solution: false -->
```python
# Test Cases
compareOutToString("git log --oneline --graph --branches --remotes", "Incorrect 'git checkout'", ["(HEAD -> main, origin/main) Added try/except"])
```

```python
!cat my_abs.py
```

<img src="./img/branches-03.png" width="400" />

## 5.3: Using topic branches

## The why and what of topic branches

When you are collaborating on a large code base, **topic branches** (or **feature branches**) can be your team's best friend.

The idea is that, when you want to develop a new feature or fix a bug, you create a branch specifically for that purpose (or "topic").  A topic branch has a very specific purpose and (ideally) touches a limited part of the code base.  It often has a limited lifetime, too; when its changes are merged into the main branch, the topic branch is usally deleted.

This has a number of technical benefits, but the greatest benefits are organizational:

* Multiple developers can work on the multiple, independent features at the same time without stepping on each others' toes
* Maintainers can more-easily understand and approve changes
* Topic branches can correspond to tasks or cards in a project management system

### `git checkout -b`: Create a new branch

Let's create a topic branch to implement some type-specific behavior in `my_abs`.  We will call this branch `types`.  We use the command `git checkout -b <branch_name>` to create a new branch at HEAD's current location.  

(*Many projects have naming conventions for topic branches. Please consult your project maintainers for details.*)

Notice that by creating the branch, we have just created a new pointer to a commit.  No files have been copied or changed, and no new commits have been made.

```python
!git checkout -b types
```

<img src="./img/branches-04.png" width="400" />

```python
!git log --oneline --graph --branches --remotes 
```

<!-- grade_id: gitcheckoutnewbranch, points: 1, grade: true, solution: false -->
```python
# Test Cases
compareOutToString("git log --oneline --graph --branches --remotes", "Incorrect 'git checkout'", ["(HEAD -> types, origin/main, main) Added try/except"])
```

### Making new changes

Now we are ready to implement our new feature.  Copy this to `my_abs.py`

``` python
import math
import numbers

def my_abs(x):
    if isinstance(x, numbers.Real):
        if x < 0:
            return -x
        else:
            return x
    elif isinstance(x, numbers.Complex):
        return math.sqrt(
            x.real ** 2 + x.imag ** 2)
    else:
        return math.nan
```

```python
!git diff
```

### Results of commit

Now let's commit our changes and observe the results:
* A new commit has been created, and both HEAD and "types" have moved forward.
* "main" and "origin/main" stay put

The effect is that your new feature is on "types", while "main" remains as a stable version.

```python
!git commit -am "Using type checks"
```

```python
!git log --oneline --graph --branches --remotes 
```

<!-- grade_id: gitnewbranchcommit, points: 1, grade: true, solution: false -->
```python
# Test Cases
compareOutToString("git log --oneline --graph --branches --remotes", "Incorrect branching", ["(HEAD -> types) Using type checks"])
```

<img src="./img/branches-05.png" width="500" />

### Making another topic branch

Now suppose we want to introduce a separate feature.  We're busy people working on multiple things.  We might have a lot of ideas, and our projects have a lot of tasks to complete.  

If we're pursuing a separate feature, we will usually create **a separate topic branch**.  And in many cases, we will want to create the topic branch **from the stable version in "main"** (instead of another topic branch).  This keeps our tasks isolated and makes task management much easier.

Let's suppose our new feature is a function to compare if two numbers are almost equal to each other.  We'll call the topic branch "almost-eq".  Our steps are:
1. Checkout "main".  If we want to branch off of "main", we need to checkout "main" first.
2. Create the topic branch "almost-eq".  



```python
!git checkout main
```

```python
!git checkout -b almost-eq
```

<img src="./img/branches-06.png" width="500" />

### Implementing `my_almost_eq`

Our implementation of `my_almost_eq` will use the existing implementation of `my_abs`.  Notice that `my_almost_eq` does not depend on the actual implementation of `my_abs`; it just depends on the input and output (the interface) of `my_abs`.  This means that someone else can make changes to `my_abs` on another topic branch, and as long as they don't change the interface, then `my_almost_eq` and any other functions that use `my_abs` will still work.

Copy this to `my_abs.py`, and note that `my_abs` is not different from "main".  

``` python
import math

def my_abs(x):
    try:
        if x < 0:
            return -x
        else:
            return x
    except TypeError:
        return math.nan

def my_almost_eq(x, y):
    return my_abs(x - y) < 1e-16
```

```python
!git diff
```

### Commtting changes to "almost-eq"

Now we are ready to commit our changes.  The effects are:
1. A new commit is created
2. HEAD and "almost-eq" move forward
3. "types", "main", and "origin/main" stay put

```python
!git commit -am "Added my_almost_eq"
```

```python
!git log --oneline --graph --branches --remotes 
```

<!-- grade_id: newbranchalmosteq, points: 1, grade: true, solution: false -->
```python
# Test Cases
compareOutToString("git log --oneline --graph --branches --remotes", "Incorrect branching", ["(HEAD -> almost-eq) Added my_almost_eq"])
```

<img src="./img/branches-07.png" width="500" />

## 5.4: Merging branches

When a topic branch is ready to be added to the main branch, it needs to **merged** into main.

In a local repo, you would use the `git merge` command.  We will use it for this demo and study how merges are resolved.  

In most collaborative workflows, merges happen on the remote repo (i.e., GitHub).  The main advantage is that GitHub has excellent support for well-defined approval processes through Pull Requests (see upcoming Module 2!).  However, on the backend, the Pull Request is ultimately resolved through a `git merge`.  And understanding how merges are resolved can be vital for making effective Pull Requests.

### Case 1: A fast-forward merge

First, suppose we like the changes in "types" and are ready to merge them into "main".  We will use the command:

```
git merge <branch_name>
```

which will merge the specified branch into the currenly checked-out branch.  Hence, to merge "types" into "main", we first need to checkout "main".  Note how HEAD moves back to "main".

```python
!git checkout main
```

```python
!git log --oneline --graph --branches --remotes 
```

<img src="./img/branches-08.png" width="500" />

Now we use `git merge types` to merge "types" into main.  In this case, "types" is a direct descendent of "main", so the merge can be resolved simply by moving "main" (and HEAD) forward to "types.  This is a case where branches very clearly act like pointers: the merge doesn't create any new commits, and it simply moves the pointer for "main" (and HEAD) forward.

This is called a **fast forward** merge and is the simplest, most unambiguous kind of merge.

```python
!git merge types
```

```python
!git log --oneline --graph --branches --remotes 
```

<!-- grade_id: ffmergebasic, points: 1, grade: true, solution: false -->
```python
# Test Cases
compareOutToString("git log --oneline --graph --branches --remotes", "fast-forward merge error", ["(HEAD -> main, types) Using type checks"])
```

<img src="./img/branches-09.png" width="500" />

### Case 2:  A merge commit

Now let's consider a case that cannot be resolved with a fast-forward merge.

Suppose that we now want to merge "almost-eq" into "main".  Unlike before, "almost-eq" is **not** a direct descendant of "main".  If Git simply moved the pointer for "main" to the same commit as "almost-eq", then all the changes from the previous merge would be left behind.

Therefore, to resolve this merge, Git will create a **new commit** with changes from both  "main" and "almost-eq".  The syntax is the same as before (`git merge almost-eq`), and git will analyze the graph of the history to determine whether to do a fast-forward or other merge.  

(*In this Jupyter Notebook, I am using the `--no-edit` flag to prevent a text editor from popping up.  This will use a default message, "Merge branch 'almost-eq'".  If you are working directly in the terminal, the `--no-edit` flag is optional.*)

```python
!git merge --no-edit almost-eq
```

```python
!git log --oneline --graph --branches --remotes 
```

<!-- grade_id: mergecommit, points: 1, grade: true, solution: false -->
```python
# Test Cases
compareOutToString("git log --oneline --graph --branches --remotes", "merge-commit error", ["(HEAD -> main) Merge branch 'almost-eq'"])
```

<img src="./img/branches-10.png" width="500" />

### Examining results of the merge

Let's look at the final results.  Now "main" contains:
* The implementation of `my_abs` with type-specific behavior (from the branch "types")
* The function `my_almost_eq` (from "almost-eq")

The merges went smoothly because our code was well-modularized and the topic branches did not touch the same parts of the code.  This is another reason why well-designed code and coordinated task management can help your collaborations go smoothly

```python
!cat my_abs.py
```

## 5.5: Branches and remotes

### Remote-tracking branches

Now it's finally time to explain "origin/main".  This is a **remote-tracking branch**, which is a local branch that indicates the state of a remote branch.  In this case, "origin/main" indicates where the branch "main" is on the remote "origin".  And indeed, if you go to your repo on GitHub, you can confirm that the state of "main" is still at the same place as "origin/main".

Let's do a `git push` and see what happens to both the remote-tracking branch and GitHub.  After the push:
* GitHub has the latest changes from main
* The remote-tracking branch "origin/main" has been moved forward to reflect the new state of the remote.

```python
!git push
```

```python
!git log --oneline --graph --branches --remotes 
```

<!-- grade_id: gitpush, points: 1, grade: true, solution: false -->
```python
# Test Cases
compareOutToString("git log --oneline --graph --branches --remotes", "git push error", ["(HEAD -> main, origin/main) Merge branch 'almost-eq'"])
```

<img src="./img/branches-11.png" width="500" />

### When the remote-tracking branch seems wrong

The very tricky thing about remote-tracking branches is:  **they are not automatically updated**.  So, if you push from one working copy, but go to another older working copy, then the remote-tracking branch on the latter will be **out-of-date**.  In the older working copy, you must explicitly **fetch** the latest changes from the remote to update the remote-tracking branch.  

Let's go back to our other working copy in `./somewhere-else/git-workflow`.  If we look at the log, we can see that:
* "main" is not up-to-date
* "origin/main" is not up-to-date either!

```python
changeDirectory(assignmentPath / "somewhere-else" / "git-workflow")
```

```python
!git log --oneline --graph --branches --remotes 
```

```python
!cat my_abs.py
```

<img src="./img/branches-12.png" width="400" />

### `git fetch`: Update the remote tracking branch

We'll get the new changes in two steps.  First, we'll use `git fetch` to update "origin/main".  This downloads new commits from the remote and moves the remote-tracking branch forward.   It does **not** change the state of the local branch "main".  

```python
!git fetch
```

```python
!git log --oneline --graph --branches --remotes 
```

<!-- grade_id: gitfetch, points: 1, grade: true, solution: false -->
```python
# Test Cases
compareOutToString("git log --oneline --graph --branches --remotes", "git fetch error", ["(origin/main, origin/HEAD) Merge branch 'almost-eq'"])
# gitHashCompare("origin/main", "origin/HEAD", "git fetch error") # Alternative method: do not use
```

<img src="./img/branches-13.png" width="500" />

### `git merge`: Merge remote-tracking branch into local branch

In many ways, remote tracking branches can be treated like other branches.  In particular, we can use them in merges.  

So, to bring "main" up-to-date with "origin/main", we will use `git merge origin/main`.  Since "origin/main" is a direct descendent of "main", this will be resolved as a fast-forward merge.

```python
!git merge origin/main
```

```python
!git log --oneline --graph --branches --remotes 
```

<!-- grade_id: gitfetchmerge, points: 1, grade: true, solution: false -->
```python
# Test Cases
compareOutToString("git log --oneline --graph --branches --remotes", "git merge error - post git fetch", ["(HEAD -> main, origin/main, origin/HEAD) Merge branch 'almost-eq'"])
#gitHashCompare("HEAD", "main", "git merge error - post git fetch") # Alternative method: do not use
#gitHashCompare("origin/main", "origin/HEAD", "git merge error - post git fetch") # Alternative method: do not use
#gitHashCompare("HEAD", "origin/HEAD", "git merge error - post git fetch") # Alternative method: do not use
```

```python
!cat my_abs.py
```

<img src="./img/branches-14.png" width="500" />

### `git pull` revisited

When we updated "main" just now, why didn't we use `git pull` like before?  We could've!  In fact, a `git pull` is just a fetch followed by a merge.  This time, I wanted to illustrate what was happening with the remote-tracking branches, so I used two explicit steps. 

The vast majority of the time, you'll be doing a `pull` instead of separate `fetch` and `merge`.  But sometimes -- like if your local history is messy -- it can be helpful to do a fetch and see what's happening before you merge. 

---

# 6: Git Commit Signing and DCO Signoff

<img src="./img/icon-gears.png" width="200" />

In professional and enterprise environments, you may encounter requirements for **signed commits** and **DCO signoff**. These are important security and legal compliance measures that help establish trust and legal clarity in collaborative projects.

## 6.1: Developer Certificate of Origin (DCO) Signoff

### What is `git commit --signoff`?

The `--signoff` flag adds a "Signed-off-by" line to your commit message, which is a **requirement in some projects** - for example, the Linux kernel and official IBM projects.

In the Linux kernel, this practice originated from lawsuits and countersuits between [SCO Group versus Linux vendors](https://en.wikipedia.org/wiki/SCO%E2%80%93Linux_disputes) and helps to solve concerns around copyright infringement. The sign-off effectively serves as a **Developer Certificate of Origin (DCO)**, asserting that:

- ✅ You wrote the code yourself, **or**
- ✅ You have the right to submit it under an appropriate open-source license, **or**
- ✅ Someone else provided it to you under such terms

### Using DCO Signoff

To sign off on a commit, use the `--signoff` or `-s` flag:

```bash
git commit --signoff -m "Fix bug in user authentication"
```

This will add a line like this to your commit message:
```
Signed-off-by: Your Name <your.email@example.com>
```

### Setting Up Automatic DCO Signoff

For projects that require DCO signoff, you can create an alias to make this automatic:

```bash
# Create an alias that always includes signoff
git config --global alias.ci 'commit --signoff'

# Now you can use:
git ci -m "Your (now signed-off) commit message"
```

---

## 6.2: Git Commit Signing (Cryptographic Signatures)

### What is Git Commit Signing?

Following GitHub's recommendations for integration branch management (e.g., `main`), many projects require commits to be **cryptographically signed**. This is **completely separate and different** from DCO signoff discussed above. Here, we mean using **GNU GPG or SSH** to digitally sign each commit with cryptographic proof.

**Important: These are NOT the same thing:**
- **DCO Signoff** (`--signoff` or `-s`): Legal attestation that you have rights to contribute the code
- **Commit Signing** (`--gpg-sign` or `-S`): Cryptographic proof that the commit came from you and hasn't been tampered with

**You can use either one, both, or neither** depending on your project's requirements.

Commit signing provides:
- **Authentication**: Cryptographically proves the commit came from you
- **Integrity**: Ensures the commit hasn't been tampered with
- **Non-repudiation**: You can't deny you made the commit
- **Verified Badge**: Shows "Verified" on GitHub commits

### Using Commit Signing

To sign a commit, you add either the `--gpg-sign` or `-S` (capital S) flag to the commit:

```bash
git commit -S -m "Add new feature with cryptographic signature"
```

### Setting Up SSH-Based Commit Signing

We'll use SSH for signing since it's simpler than GPG setup. For full instructions including GPG, see [GitHub's documentation](https://docs.github.com/en/authentication/managing-commit-signature-verification).

#### Step 1: Create an SSH Key (if you don't have one)

```bash
ssh-keygen -t ed25519 -C "your_email@example.com"
```

#### Step 2: Add Signing Key to GitHub

1. Go to your GitHub account's **Settings** → **SSH and GPG keys**
2. Click the green **"New SSH key"** button
3. Select **"Signing key"** in the key type menu (**Important**: Choose "Signing key", not "Authentication key")
4. Name the key and paste the contents of your public key file (`~/.ssh/id_ed25519.pub`)

#### Step 3: Configure Local Git Environment

```bash
# Tell Git to use SSH for signing
git config --global gpg.format ssh

# Specify your signing key
git config --global user.signingkey ~/.ssh/id_ed25519.pub

# Enable automatic signing for all commits
git config --global commit.gpgsign true
```

### The Goal: Simple `git commit -m` Commands

**After proper setup, you should only need to use:**

```bash
git commit -m "Your commit message"
```

**That's it!** No extra flags needed. Here's what happens automatically based on your configuration:

- If you set `commit.gpgsign true` → **Automatically signed**
- If you need DCO signoff → Set up an alias (see below)
- If you need both → Set up an alias for both (see below)

### One Command to Rule Them All: Smart Alias Setup

**The simplest approach is to set up your `git commit` command to automatically include both DCO signoff and GPG signing.** This way, you're covered for any project requirements:

```bash
# Override the default commit command to include both signoff and signing
git config --global alias.commit 'commit --gpg-sign --signoff'
```

After setting this alias, **every time you use `git commit -m "message"`, it will automatically:**
- Add DCO signoff (`Signed-off-by: Your Name <email>`)
- Cryptographically sign the commit (shows "Verified" on GitHub)

### Why This Approach Works Best

**DCO signoff is harmless on projects that don't require it** - it's just an extra line in your commit message that proves you wrote the code. Think of it as additional documentation that you have the rights to contribute the code. Projects that don't require DCO will simply ignore the `Signed-off-by` line.

**Benefits of using both by default:**
- **One setup, works everywhere** - No need to remember different commands for different projects
- **Maximum compliance** - Covers both legal (DCO) and technical (signing) requirements
- **Better documentation** - Shows clear authorship and legal attestation
- **Simple workflow** - Just use `git commit -m "message"` everywhere

### Alternative: Project-Specific Aliases

If you prefer more granular control, you can set up specific aliases:

```bash
# For DCO signoff only
git config --global alias.dco 'commit --signoff'

# For signing only  
git config --global alias.sign 'commit --gpg-sign'

# For both (recommended)
git config --global alias.secure 'commit --gpg-sign --signoff'
```

Then use:
```bash
git dco -m "Your message"     # DCO only
git sign -m "Your message"    # Signing only
git secure -m "Your message"  # Both
```

### Recommended Setup Commands

**For the "one command for everything" approach (recommended):**

```bash
# Set up automatic signing
git config --global gpg.format ssh
git config --global user.signingkey ~/.ssh/id_ed25519.pub
git config --global commit.gpgsign true

# Override commit to always include DCO signoff
git config --global alias.commit 'commit --signoff'
```

Now `git commit -m "Your message"` will automatically include both DCO signoff and cryptographic signing!

### Important Notes

1. **Automatic Signing**: Setting `commit.gpgsign true` means **all commits are automatically signed**. You don't need to remember any flags.

2. **DCO Signoff is Always Safe**: Adding DCO signoff to all commits doesn't hurt - it's just extra documentation that you have rights to contribute the code.

3. **Repository-Specific Settings**: If you don't want to make these changes globally, omit the `--global` flag and run these commands in each repo where they should apply.

4. **Fixing Commits**: If you need to fix a commit *that you haven't pushed yet* (recommended for unpublished changes only):
   ```bash
    git commit --amend -m "Updated message"  # Uses your existing config
   ```
   > **Note:** Only use `git commit --amend` for commits that have not been pushed to a shared repository. Amending published commits can cause problems for collaborators.

### Verifying Signed Commits

You can verify signed commits using:

```bash
# Show signature information in log
git log --show-signature

# Verify specific commit
git verify-commit <commit-hash>

# Show one-line log with signature status
git log --oneline --show-signature
```

On GitHub, signed commits will show a **"Verified"** badge next to the commit, indicating cryptographic verification has passed.